---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## First step to awesomeness

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise. In fact it's an awesome example exercise!! locked? this is pepito, while locked.

`@instructions`
It should be clear at this point that the final objective of this course is awesomeness.

With this in mind, just go ahead and print the word "awesome".

`@hint`
It might have something to do with the `print` function...

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
print('awesome')
```

`@sct`
```{python}
Ex().has_output("awesome", pattern = False)
success_msg('Congratulations! Your path to awesomeness has began!')
```
