---
title: chap2
description: de
---

## Insert exercise title here

```yaml
type: NormalExercise
key: df9f38eca3
xp: 100
```

<!-- Coding exercise documentation: https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises -->
blahhhhbkj kjhkl hkljh kljhlkj hlk hkl hlk

`@instructions`


`@hint`
<!-- Hint best practices: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices -->
Add a hint to support the learners when they get stuck.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
