---
title: chap2
description: de
---

## chap2, ex1

```yaml
type: NormalExercise
key: df9f38eca3
xp: 100
```

<!-- Coding exercise documentation: https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises -->
blahhhhbkj kjhkl hkljh kljhlkj hlk hkl hlk foo

`@instructions`


`@hint`
<!-- Hint best practices: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices -->
Add a hint to support the learners when they get stuck.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## the title of ex

```yaml
type: VideoExercise
key: 3579ec3b84
xp: 50
```

`@projector_key`
50d1069640da0172505ff26fdd6328f8

---

## the coding exercise

```yaml
type: NormalExercise
key: f67ad7dd3a
xp: 100
```

<!-- Guidelines for contexts: https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->

`@instructions`
<!-- Guidelines for instructions https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->
- Instruction 1
- Instruction 2
- foo

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
```
